
# Module 1 Assignment

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says # YOUR CODE HERE. Do not write your answer anywhere else other than where it says # YOUR CODE HERE. Anything you write elsewhere will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Kernel, and restart the kernel and run all cells (Restart & Run all).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select File → Save and CheckPoint).

5. All work must be your own. If you do use any code from another source (such as a course notebook or a website) you need to properly cite the source.


```python
from nose.tools import assert_equal, assert_almost_equals, assert_not_equal
import sys
import numpy as np
import warnings
warnings.filterwarnings("ignore", category=DeprecationWarning) 
```

# Problem 1: Running your first Python code

- Type "x = 3.14" without quotation marks in the below cell and run it.

<!-- grade_id: p1-answer, points: 0, grade: false, solution: true -->
```python

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: p1-test, points: 1, grade: true, solution: false -->
```python
try:
    if ((x==3.14) & (type(x)==float)):
        print('Great! You ran your first Python code!')
    else:
        print('x is not the right value check your code and run again')
except NameError:
    print('x is not defined! You need to run the code cell!')

assert_equal(x,3.14,msg='x is not equa lto 3.14!')
```

# Problem 2: Use cell magic to view a plot inline

- In the below code cell, type "%matplotlib inline" and then run the cell. This allows Python plots to be shown inside the notebook.

<!-- grade_id: p2-answer, points: 0, grade: false, solution: true -->
```python

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: p2-test, points: 1, grade: true, solution: false -->
```python
theta = np.arange(0., np.pi, 0.01)
y = np.sin(theta)

import matplotlib.pyplot as plt

plt.plot(theta, y)
plt.xlabel(r"$\theta$")
plt.ylabel(r"$\sin$($\theta$)")
plt.title("My Plot")
plt.show()
```

# Problem 3: Change the cell type

- The below cell is in markdown format. Change the cell type so that the cell type is now "code" and then run the cell. Note that although "#YOUR CODE HERE" appears in the cell, no code writing is necessary. 


- HINT: Look at the "Cell" tab in the notebook menu bar.

#YOUR CODE HERE

z = 'Change my cell type to code so I can run!'

<!-- grade_id: p3-test, points: 1, grade: true, solution: false -->
```python
try:
    if ((z=='Change my cell type to code so I can run!') & (type(z)==str)):
        print('Great! You changed the cell type')
except NameError:
    print('z is not defined! Check the cell type!')
    
assert_equal(type(z),str,msg='z is not defined! change the cell type!')
```

# Problem 4: Write a Markdown header

- The below cell is in code format, but we can use the IPython.display module to write Markdown formatted code as Python strings, and then display it as markdown.


- In the below cell, in between the two ` '''` statements, write markdown code that would produce the **largest possible** markdown header. The header text should be "My first header".

<!-- grade_id: p4-answer, points: 0, grade: false, solution: true -->
```python
md='''

# YOUR CODE HERE
raise NotImplementedError()


'''
```

```python
#Let's view the markdown
from IPython.display import Markdown
display(Markdown(md))
```

<!-- grade_id: p4-test, points: 1, grade: true, solution: false -->
```python
text = md.find('# My first header')
assert_not_equal(text,-1,msg='Make sure you are using the largest header.')
```

-----

**&copy; 2019: Gies College of Business at the University of Illinois.**

This notebook is released under the [Creative Commons license CC BY-NC-SA 4.0][ll]. Any reproduction, adaptation, distribution, dissemination or making available of this notebook for commercial use is not allowed unless authorized in writing by the copyright holder.

[ll]: https://creativecommons.org/licenses/by-nc-sa/4.0/legalcode