# Module 1 Assignment


A few things you should keep in mind when working on assignments:

1. Run the first code cell to import modules needed by this assignment before proceeding to problems.
2. Make sure you fill in any place that says `# YOUR CODE HERE`. Do not write your answer anywhere else other than where it says `# YOUR CODE HERE`. Anything you write elsewhere will be removed or overwritten by the autograder.
3. Each problem has an autograder cell below the answer cell. Run the autograder cell to check your answer. If there's anything wrong in your answer, the autograder cell will display error messages.
4. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Kernel, and Restart & Run all. If the notebook runs through the last code cell without an error message, you've answered all problems correctly.
5. Make sure that you save your work (in the menubar, select File → Save and CheckPoint).

-----

# Run Me First!

```python
from nose.tools import assert_equal, assert_almost_equals, assert_not_equal
import sys
import math
import warnings
warnings.filterwarnings("ignore", category=DeprecationWarning) 
```

# Problem 1: Creating a Python string

- Type `x = 'This is my first string'` in the below cell and run it.

<!-- grade_id: p1-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: p1-test, points: 1, grade: true, solution: false -->
```python
try:
    if ((x=='This is my first string') & (type(x)==str)):
        print('Great! You created a string!')
    else:
        print('x is not the right value check your code and run again')
except NameError:
    print('x is not defined! You need to run the code cell!')
    
assert_equal(type(x),str, msg='x is not a string!')
```

# Problem 2: Import numpy

- Type `import numpy` in the below cell to import the numpy module.


<!-- grade_id: p2-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

```

<!-- grade_id: p2-test, points: 1, grade: true, solution: false -->
```python
if 'numpy' in dir():
    print('numpy module has been successfully imported')
else:
    print('numpy module has not been imported')
    
assert_equal(numpy.array([1])[0],1,msg='Make sure you ran the code cell to import numpy!')
```

# Problem 3: Add emphasis to Markdown text

- The below cell is in code format, but we can use the IPython.display module to write Markdown formatted code as Python strings, and then display it as markdown.


- In the below cell, in between the two ` '''` statements, write markdown code that does the following

    - Writes the word "bold" in **bold**
    
    - Writes the word "italics" in *italics*

- NOTE: When writing your markdown code, dont put any spaces between the Markdown characters that initialize the bold/italicize environments, and your text within those environments.

After you run your answer code cell and the cell bolow, the output of the cell below the answer cell should look like this:

# YOUR CODE HERE

**bold**  
*italics*  

<!-- grade_id: p3-answer, points: 0, grade: false, solution: true -->
```python
md='''

# YOUR CODE HERE

'''
```

```python
#Let's view the markdown
from IPython.display import Markdown
display(Markdown(md))
```

<!-- grade_id: p3-test, points: 1, grade: true, solution: false -->
```python
text = md.find('**bold')
if text == -1:
    text = md.find('__bold')
assert_not_equal(text,-1,msg='You did not bold the word bold.')
text = md.find('*ital')
if text == -1:
    text = md.find('_ital')
assert_not_equal(text,-1,msg='You did not italicize the word italics')
```

# Problem 4: Create a Markdown list

- The below cell is in code format, but we can use the IPython.display module to write Markdown formatted code as Python strings, and then display it as markdown.


- In the below cell, in between the two ` '''` statements, write markdown code that creates a markdown list with 3 items. The items should be numbered 1 through 3, respectively, and there should be a `.` character after each number. Name each item "Item 1", "Item 2", and "Item 3", respectively. Note that the "I" in Item sould be capitalized.

After you run your answer code cell and the cell bolow, the output of the cell below the answer cell should look like this:
# YOUR CODE HERE  

1. Item 1
2. Item 2
3. Item 3


<!-- grade_id: p4-answer, points: 0, grade: false, solution: true -->
```python
md='''

# YOUR CODE HERE


'''
```

```python
#Let's view the markdown
from IPython.display import Markdown
display(Markdown(md))
```

<!-- grade_id: p4-test, points: 1, grade: true, solution: false -->
```python
text = md.find('1.')
assert_not_equal(text,-1,msg='Number 1. is mising from your markdown list')
text = md.find('Item')
assert_not_equal(text,-1,msg='Make sure you named each item in the list properly.')
```

# Problem 5: Show code in Markdown

- The below cell is in code format, but we can use the IPython.display module to write Markdown formatted code as Python strings, and then display it as markdown.


- In the below cell, in between the two ` '''` statements, write markdown code that renders Python code. The Python code to be rendered is the statement `x = 'this is a python string'`. 

- NOTE: When writing your markdown code, dont put any spaces between the Markdown characters that initialize the "code" environment.

After you run your answer code cell and the cell bolow, the output of the cell below the answer cell should look like this:

# YOUR CODE HERE  

`x = 'this is a python string'`

<!-- grade_id: p5-answer, points: 0, grade: false, solution: true -->
```python
md='''

# YOUR CODE HERE


'''
```

```python
#Let's view the markdown
from IPython.display import Markdown
display(Markdown(md))
```

<!-- grade_id: p5-test, points: 1, grade: true, solution: false -->
```python
text = md.find('`x')
assert_not_equal(text,-1,msg='Make sure you are using the right characters to render the code.')
text = md.find('python')
assert_not_equal(text,-1,msg='Make sure you type the string correctly.')
```