
# Module 2 Assignment

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says # YOUR CODE HERE. Do not write your answer anywhere else other than where it says # YOUR CODE HERE. Anything you write elsewhere will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Kernel, and restart the kernel and run all cells (Restart & Run all).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select File → Save and CheckPoint).

5. All work must be your own, if you do use any code from another source (such as a course notebook or a website) you need to properly cite the source.


```python
from nose.tools import assert_equal, assert_almost_equals, assert_true
import sys
import math

import warnings
warnings.filterwarnings("ignore", category=DeprecationWarning) 
```

# Problem 1: Import numpy

For this problem, you will import the Python module numpy.

- Import the Python module numpy.

<!-- grade_id: p1_ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: p1_test, points: 1, grade: true, solution: false -->
```python
if 'numpy' in dir():
    print('numpy module has been successfully imported')
else:
    print('numpy module has not been imported')

assert_true("numpy" in dir(), msg="You need to import the numpy library")
```

# Problem 2: Create a variable

Create a variable named `a` which is equal to the number 3

<!-- grade_id: p2_ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: p2_test, points: 1, grade: true, solution: false -->
```python
try:
    if (a == 3):
        print('Great! a is equal to 3.')
    else:
        raise Exception('a is defined but is not the right value.')
except NameError:
    raise Exception('a is not defined! This means it has not been saved as a variable.')

b = 390877006.48625034  # You will need this variable for the next problem.
```

# Problem 3: Calculate $7(a+b)-\pi b$

- Use Python math operations and the Python math module to calculate the above expression. Save the result in a variable called `d`.

<!-- grade_id: p3_ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: p3_test, points: 1, grade: true, solution: false -->
```python
try:
    assert_almost_equals(d,1508162734.3693783,msg='d is not the correct value')
except NameError:
    raise Exception('d is not defined! This means you did not assign a value to the variable d.')

```

# Problem 4: Create a function to determine if a number is even
- Write a function called `is_even_number` that accepts one parameters called `num`.   
   
- Inside of this function, return True if num is even number, otherwise return False.

- Hint: Use remainder(%) operator


<!-- grade_id: p4_ans, points: 0, grade: false, solution: true -->
```python
def is_even_number(num):
    '''
    Returns True if num is even number, otherwise False
    
    Parameters
    ----------
    num : an integer
    
    Returns
    -------
    True if num is even number, otherwise False
    '''
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: p4_test, points: 1, grade: true, solution: false -->
```python
assert_true(is_even_number(2), msg='2 is an even number.')
assert_true(not is_even_number(1), msg='1 is not an even number')
```

-----

**&copy; 2019: Gies College of Business at the University of Illinois.**

This notebook is released under the [Creative Commons license CC BY-NC-SA 4.0][ll]. Any reproduction, adaptation, distribution, dissemination or making available of this notebook for commercial use is not allowed unless authorized in writing by the copyright holder.

[ll]: https://creativecommons.org/licenses/by-nc-sa/4.0/legalcode