# Module 2 Assignment

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says # YOUR CODE HERE. Do not write your answer anywhere else other than where it says # YOUR CODE HERE. Anything you write elsewhere will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Kernel, and restart the kernel and run all cells (Restart & Run all).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select File → Save and CheckPoint).

5. All work must be your own, if you do use any code from another source (such as a course notebook or a website) you need to properly cite the source.

```python
from nose.tools import assert_equal, assert_true, assert_almost_equals
import sys
import warnings
warnings.filterwarnings("ignore", category=DeprecationWarning) 
```

# Problem 1: Import math

For this problem, you will import the Python module math.

- Import the Python module math.

<!-- grade_id: p1_ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: p1_test, points: 1, grade: true, solution: false -->
```python
if 'math' in dir():
    print('math module has been successfully imported')
else:
    print('math module has not been imported')

assert_true("math" in dir(), msg="You need to import the math library")
```

# Problem 2: Calculate $3cos(\pi + 4) - 12sin(\pi^{3})^{-1}$

- Use Python math operations, and the Python math module, to calculate the above expression. Save the result in a variable called "my_calculation".

<!-- grade_id: p2_ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: p2_test, points: 1, grade: true, solution: false -->
```python
try:
    assert_almost_equals(my_calculation,32.089869248682895,msg='my_calculation is not the correct value')
except NameError:
    print('my_calculation is not defined! This means it has not been saved as a variable.')

print(my_calculation)
```

# Problem 3: Adding remainders

- Calculate the remainder of 7 divided by 6,    
   
   
- and calculate the remainder of 243 divided by 192. 


- Add both of these values together and save the result in a variable called "my_remainder".


<!-- grade_id: p3_ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: p3_test, points: 1, grade: true, solution: false -->
```python
try:
    assert_almost_equals(my_remainder,52,msg='my_remainder is not the correct value')
except NameError:
    print('my_remainder is not defined! This means it has not been saved as a variable.')

print(my_remainder)
```

# Problem 4: Create a function to determine which value is greater  

- Write a function called find_greater that accepts 2 parameters called a and b.


- Inside of the function use an if statement to determine if a is greater than b.

    - If the value in a is greater than the value in b return variable a.
    
    - If the value in b is greater than the value in a return variable b.


<!-- grade_id: p4_ans, points: 0, grade: false, solution: true -->
```python
def find_greater(a, b):
    '''
    Returns the highest value between a and b.
    
    Parameters
    ----------
    a : an integer or floating point number
    b : an integer or floating point number
    
    Returns
    -------
    a floating point number that has the highest value
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: p4_test, points: 1, grade: true, solution: false -->
```python
assert_equal(find_greater(3, 4), 4, msg='The greater value between 3 and 4 is 4.')
assert_equal(find_greater(4, 3), 4, msg='The greater value between 4 and 3 is 4.')
assert_equal(find_greater(3.1459, 3.14), 3.1459, msg='The greater value between 3.1459 and 3.14 is 3.1459')

```

# Problem 5: Create a function to determine which value is the greatest

- Write a function called find_greatest that accepts 3 parameters called a, b, and c.


- Inside of the function use an if, else, and/or elif statements to determine which variable between a, b, and c has the highest value. Return this value.

**HINT: There's more than 1 way to solve this problem.**

<!-- grade_id: p5_ans, points: 0, grade: false, solution: true -->
```python
def find_greatest(a, b, c):
    '''
    Returns the highest value between a, b, and c.
    
    Parameters
    ----------
    a : an integer or floating point number
    b : an integer or floating point number
    c : an integer or floating point number

    Returns
    -------
    a floating point number that has the highest value
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: p5_test, points: 1, grade: true, solution: false -->
```python
assert_equal(find_greatest(2, 3, 4), 4, msg='The greater value between 2, 3 and 4 is 4.')
assert_equal(find_greatest(3, 2, 4), 4, msg='The greater value between 3, 2 and 4 is 4.')
assert_equal(find_greatest(4, 3, 2), 4, msg='The greater value between 4, 3 and 2 is 4.')

```

-----

**&copy; 2019: Gies College of Business at the University of Illinois.**

This notebook is released under the [Creative Commons license CC BY-NC-SA 4.0][ll]. Any reproduction, adaptation, distribution, dissemination or making available of this notebook for commercial use is not allowed unless authorized in writing by the copyright holder.

[ll]: https://creativecommons.org/licenses/by-nc-sa/4.0/legalcode