
# Module 1 Assignment(Demo)

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says # YOUR CODE HERE. Do not write your answer anywhere else other than where it says # YOUR CODE HERE. Anything you write elsewhere will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Kernel, and restart the kernel and run all cells (Restart & Run all).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select File → Save and CheckPoint).

5. All work must be your own. If you do use any code from another source (such as a course notebook or a website) you need to properly cite the source.


```python
from nose.tools import assert_equal, assert_almost_equals, assert_not_equal
import sys
import numpy as np
import warnings
warnings.filterwarnings("ignore", category=DeprecationWarning) 
```

# Problem 1: Running your first Python code

- Assign value 1 to variable a

<!-- grade_id: p1-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
a = 1
### END SOLUTION
```

<!-- grade_id: p1-test, points: 1, grade: true, solution: false -->
```python
assert_equal(a,1,msg='a is not equal to 1!')
```

# Problem 2: Running your second Python code

- Assign value 2 to variable b

<!-- grade_id: p2-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
b = 2
### END SOLUTION
```

<!-- grade_id: p2-test, points: 1, grade: true, solution: false -->
```python
assert_equal(b,2,msg='b is not equal to 2!')
```

-----

**&copy; 2019: Gies College of Business at the University of Illinois.**

This notebook is released under the [Creative Commons license CC BY-NC-SA 4.0][ll]. Any reproduction, adaptation, distribution, dissemination or making available of this notebook for commercial use is not allowed unless authorized in writing by the copyright holder.

[ll]: https://creativecommons.org/licenses/by-nc-sa/4.0/legalcode