
# Demo Assignment

A few things you should keep in mind when working on assignments:

1. Run the first code cell to import modules needed by this assignment before proceeding to problems.
2. Make sure you fill in any place that says `# YOUR CODE HERE`. Do not write your answer anywhere else other than where it says `# YOUR CODE HERE`. Anything you write elsewhere will be removed or overwritten by the autograder.
3. Each problem has an autograder cell below the answer cell. Run the autograder cell to check your answer. If there's anything wrong in your answer, the autograder cell will display error messages.
4. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Kernel, and Restart & Run all. If the notebook runs through the last code cell without error message, you've answered all problems correctly.
5. Make sure that you save your work (in the menubar, select File → Save and CheckPoint).

-----


# Run Me First!

```python
from nose.tools import assert_equal, assert_almost_equals, assert_not_equal
import sys
import numpy as np
import warnings
warnings.filterwarnings("ignore", category=DeprecationWarning) 
```

# Problem 1: Running your first Python code

- Assign value 1 to variable a

<!-- grade_id: p1-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
a = 1
### END SOLUTION
```

<!-- grade_id: p1-test, points: 1, grade: true, solution: false -->
```python
assert_equal(a,1,msg='a is not equal to 1!')
```

# Problem 2: Running your second Python code

- Assign value 2 to variable b

<!-- grade_id: p2-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
b = 2
### END SOLUTION
```

<!-- grade_id: p2-test, points: 1, grade: true, solution: false -->
```python
assert_equal(b,2,msg='b is not equal to 2!')
```