# Module 4 Assignment

A few things you should keep in mind when working on assignments:

1. Make sure you fill in any place that says `YOUR CODE HERE`. Do **not** write your answer anywhere else other than where it says `YOUR CODE HERE`. Anything you write elsewhere will be removed or overwritten by the autograder.

2. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select _Kernel_, and restart the kernel and run all cells (_Restart & Run all_).

3. Do not change the title (i.e. file name) of this notebook.

4. Make sure that you save your work (in the menubar, select _File_ → _Save and CheckPoint_).


```python
from nose.tools import assert_equal, assert_almost_equal, assert_true
import sys
import os
import pandas as pd
```

-----

# Problem 1: Importing the NumPy module

In the code cell below import the NumPy module using the alias `np`. If done incorrectly you may fail other assert statements.

-----

<!-- grade_id: p1_ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: p1_test, points: 1, grade: true, solution: false -->
```python
assert_true('np' in  dir(), msg='import the numpy module using the alias np.')
```

-----

# Problem 2: Creating a NumPy array

In the code cell below the function `gen_data` accepts one parameter `n` which is an integer.

For this problem:

- Create a NumPy array containing integers from 0 to n incremented by 1.

- The final entry in the array should be n - 1.
  
- Return this newly created array.

-----

<!-- grade_id: p2_ans, points: 0, grade: false, solution: true -->
```python
def gen_data(n):
    '''
    Generates a NumPy array of integers 
    from 0 to n - 1
    
    Parameters
    ----------
    n: integer value indicating the end of 
    the interval
    
    Returns
    -------
    returns the NumPy array
    
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: p2_test, points: 1, grade: true, solution: false -->
```python
ans_10obs = gen_data(10)
assert_equal(type(ans_10obs), type(np.array(0)))
assert_equal(ans_10obs.tolist(), 
             [x for x in range(10)])

ans_120obs = gen_data(120)
assert_equal(type(ans_120obs), type(np.array(0)))
assert_equal(ans_120obs.tolist(), 
             [x for x in range(120)])
```

-----

# Problem 3: NumPy Summary Functions

In the code cell below the function `median` accepts one parameter `a` which is a NumPy array.

For this problem:
- Use the NumPy `median` function to compute the median value.
- Return this value.

-----

<!-- grade_id: p3_ans, points: 0, grade: false, solution: true -->
```python
def median(a):
    '''
    Returns the median value in a NumPy array
    
    Parameters
    ----------
    a: NumPy array
    
    Returns
    -------
    the median value of the array
    '''
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: p3_test, points: 1, grade: true, solution: false -->
```python
ans3_1 = median(ans_10obs)
assert_equal(ans3_1, 4.5, msg="Your answer does not match the solutions")

ans3_2 = median(ans_120obs)
assert_equal(ans3_2, 59.5, msg="Your answer does not match the solutions")
```

-----

# Problem 4: Select one columns from a DataFrame

The code cell below reads a DataFrame from a file. Please select one column that contains **sepal length** data as a pandas **Series** from the DataFrame and assign it to variable **sp_length**.

-----

```python
df = pd.read_csv('data/iris.csv')
df.head()
```

<!-- grade_id: p4-ans, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: p4-test, points: 1, grade: true, solution: false -->
```python
assert_equal(type(sp_length), pd.core.series.Series, msg="Your answer does not match the solutions")
assert_equal(sp_length.shape[0], 150, msg="Your answer does not match the solutions")
assert_equal(sp_length[0], 5.1, msg="Your answer does not match the solutions")
assert_equal(sp_length[1], 4.9, msg="Your answer does not match the solutions")

```

-----

**&copy; 2019: Gies College of Business at the University of Illinois.**

This notebook is released under the [Creative Commons license CC BY-NC-SA 4.0][ll]. Any reproduction, adaptation, distribution, dissemination or making available of this notebook for commercial use is not allowed unless authorized in writing by the copyright holder.

[ll]: https://creativecommons.org/licenses/by-nc-sa/4.0/legalcode