# Module 3 Assignment


A few things you should keep in mind when working on assignments:

1. Run the first code cell to import modules needed by this assignment before proceeding to problems.
2. Make sure you fill in any place that says `# YOUR CODE HERE`. Do not write your answer anywhere else other than where it says `# YOUR CODE HERE`. Anything you write elsewhere will be removed or overwritten by the autograder.
3. Each problem has an autograder cell below the answer cell. Run the autograder cell to check your answer. If there's anything wrong in your answer, the autograder cell will display error messages.
4. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Kernel, and Restart & Run all. If the notebook runs through the last code cell without error message, you've answered all problems correctly.
5. Make sure that you save your work (in the menubar, select File → Save and CheckPoint).

-----

# Run Me First!

```python
import numpy as np
from nose.tools import assert_equal, assert_is_instance

import warnings
warnings.filterwarnings('ignore')
```

# Problem 1:  BooHoo in a list

- Write a function called "BooHoo" which takes in an integer, `n`, and stores all integers from 1 to `n` in a list. However, if an integer in the list is divisible by 3, the list should instead contain "Boo". If the integer is divisible by 5 the list should instead contain "Hoo". If the integer is divisible by both 3 and 5, the list should instead contain "BooHoo".


- As an example, if the input is 5, the output of the function should be a list  
`[1, 2, 'Boo', 4, 'Hoo', 'Boo', 7, 8, 'Boo', 'Hoo', 11, 'Boo', 13, 14, 'BooHoo']`.

<!-- grade_id: p1-answer, points: 0, grade: false, solution: true -->
```python
def BooHoo(n):
    '''
    A function which returns a list of integers, and replaces integers divisible by 3 with "Boo" and integers
    divisible by 5 with "Hoo" and integers divisible by 3 and 5 with "BooHoo"
    
    
    Parameters
    -------
    n: an integer
    
    
    Returns
    --------

    final_list: a Python list
    
    '''
    
    # YOUR CODE HERE
    
    
    return final_list
```

<!-- grade_id: p1-test, points: 1, grade: true, solution: false -->
```python
test_boo = BooHoo(15)
print(test_boo)
assert_is_instance(test_boo,list,msg='Your function did not return a list')
assert_equal(len(test_boo),15,msg='Your output is not the correct length')
assert_equal(test_boo[-1],'BooHoo',msg='The function did not return the right output')
assert_equal(test_boo[0],1,msg='The function did not return the right output')
assert_equal(test_boo[2],'Boo',msg='The function did not return the right output')
```

# Problem 2: Square elements if they are divisible by 3.


- Write a function called "square_threes" that takes in an integer, `n`, checks all integers from 1 to `n`, and saves that integer squared in a list if that integer is divisible by 3.


- As an example, if `n` is equal to 6, the function should return `[9,36]`.


- Note that the returned list should *not* contain 0(since number starts from 1).


<!-- grade_id: p2-answer, points: 0, grade: false, solution: true -->
```python
def square_threes(n):
    '''
    A function which takes in an integer and returns a list containing the squared  integer from 1 to n, if
    that integer is divisible by 3
    
    
    Parameters
    -------
    
    n: an integer
    
    
    Returns
    --------
    
    output_list: a Python list containing the squared elements
    
    '''
    
    
    # YOUR CODE HERE
    
    
    return output_list
```

<!-- grade_id: p2-test, points: 1, grade: true, solution: false -->
```python
squared = square_threes(25)
print(squared)
assert_equal(type(squared), list, msg ='Your function did not return a list')
assert_equal(squared[0], 9, msg ='The first element is incorrect')
assert_equal(squared[-1], 576, msg = 'The last element is incorrect')
```

# Problem 3: Summing a list of lists.

- Write a function called "sum_of_list_of_lists" which takes in a Python list of lists and returns a single sum of all the elements in the list of lists. The output should be a single number.

- Assign the sum to variable sum_of_lists and return it from the function.

- Note that the elements of input_list are lists of integers, for example:  
`[[1,2,3], [5,6], [4]]`

<!-- grade_id: p3-answer, points: 0, grade: false, solution: true -->
```python
def sum_of_list_of_lists(input_list):
    '''
    A function which takes in a Python list of lists and sums up all their elements
    
    
    Parameters
    -------
    
    input_list: a Python list of lists
    
    
    Returns
    --------
    
    sum_of_lists: the aggregate sum of all elements in the list of lists
    '''
    # YOUR CODE HERE
    
    
    return sum_of_lists
```

<!-- grade_id: p3-test, points: 1, grade: true, solution: false -->
```python
list_of_lists = [[1, 4, 54], [23, 45], [64, 11, 23], [53, 2, 6, 45]] 
added_lists = sum_of_list_of_lists(list_of_lists)
assert_equal(type(added_lists), int, msg ='Your function did not return an integer')
assert_equal(added_lists, 331, msg='Your function did not return the right value')
```

# Problem 4: Get reversed list

- Write a function called "get_reversed_list" which takes in a list, and then returns a reversed list.
- As an example, if input_list is [1,2,3], the function should return a list [3,2,1].


<!-- grade_id: p4-answer, points: 0, grade: false, solution: true -->
```python
def get_reversed_list(input_list):
    '''
    A function which takes a list and returns a reversed list.
    
    
    Parameters
    -------
    input_list: a Python list
        
    
    Returns
    --------
    
    reversed_list: a reversed Python list
    
    '''
    
    # YOUR CODE HERE
    
    return reversed_list
```

<!-- grade_id: p4-test, points: 1, grade: true, solution: false -->
```python
input_list = [24,3451,3412,64,212]
reversed_list = get_reversed_list(input_list)
print(reversed_list)
assert_equal(type(reversed_list), list, msg ='Your function did not return a list')
assert_equal(reversed_list[0], 212, msg ='The returned list is incorrect')
assert_equal(reversed_list[-1], 24, msg = 'The returned list is incorrect')
```

# Problem 5: Concatenating strings

- Write a function called "add_strings" which takes in 2 python strings and concatenates them. The function should put a space in between the two strings.
- As an example, if string1="Love" and string2="python", the returned string should be "Love python"


<!-- grade_id: p5-answer, points: 0, grade: false, solution: true -->
```python
def add_strings(string1,string2):
    '''
    A function which takes in two strings and concatenates them together
    
    
    Parameters
    -------
    
    string1: a Python string
    
    string2: a Python string
        
    
    Returns
    --------
    
    final_string: a Python string which is a concatenation of string1 and string2    
    '''
    
    
    # YOUR CODE HERE
    
    
    return final_string
```

<!-- grade_id: p5-test, points: 1, grade: true, solution: false -->
```python
final_string = add_strings('I love','Python')
print(final_string)
assert_equal(len(final_string), 13,msg='Your string is the wrong length')
assert_equal(final_string[:10], 'I love Pyt',msg='Your function returned an incorrect value')
```