# Module 6 Assignment


A few things you should keep in mind when working on assignments:

1. Run the first code cell to import modules needed by this assignment before proceeding to problems.
2. Make sure you fill in any place that says `# YOUR CODE HERE`. Do not write your answer anywhere else other than where it says `# YOUR CODE HERE`. Anything you write elsewhere will be removed or overwritten by the autograder.
3. Each problem has an autograder cell below the answer cell. Run the autograder cell to check your answer. If there's anything wrong in your answer, the autograder cell will display error messages.
4. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Kernel, and Restart & Run all. If the notebook runs through the last code cell without an error message, you've answered all problems correctly.
5. Make sure that you save your work (in the menubar, select File → Save and CheckPoint).

-----

# Run Me First!

```python
import pandas as pd
import numpy as np
import matplotlib as mpl
import matplotlib.pyplot as plt
from nose.tools import assert_equal, assert_almost_equal, assert_is_instance, assert_is_not
import seaborn as sns
sns.set()
import warnings
warnings.filterwarnings('ignore')
%matplotlib inline

```

```python
#Make sure the Dow Jones data is loaded in
df = pd.read_csv('./data/dow_jones_index.data')
df.head()
```

# Problem 1: Create a boxplot of the opening stock price for a given stock.

In this problem you will create a boxplot to view the distribution of the opening price for a given stock from the Dow Jones dataset.

- Write a function called "open_boxplot" that takes in a DataFrame and the ticker name of a stock and outputs a boxplot of that stocks opening price.


- Use the seaborn module to create the boxplot.


- Give the plot a descriptive title.


- Give the plot a descriptive x-axis label.


- Do *NOT* give the plot a y-axis label.

<!-- grade_id: p1-answer, points: 0, grade: false, solution: true -->
```python
def open_boxplot(df,ticker_name):
    '''
    A function which creates a boxplot of the opening price from the Dow Jones data.
        
    Parameters
    -------
    df: a pandas dataframe containing stock price data from the Dow Jones dataset
    
    ticker_name: the ticker name of the stock of interest
        
    Returns
    --------
    
    ax: a matplotlib axes object containing a boxplot of the stock's opening prices from the Dow Jones dataset    
    '''

    # YOUR CODE HERE
    
    
    return ax
```

Below we again look at the stock for Alcoa Corp (ticker name: AA)

```python
my_boxplot = open_boxplot(df,'AA')
plt.show()
plt.close()
```

<!-- grade_id: p1-test, points: 1, grade: true, solution: false -->
```python
assert_is_instance(my_boxplot, mpl.axes.Axes, msg='Return a matplotlib axes object.')  
assert_is_not(len(my_boxplot.title.get_text()), 0, msg="Your plot doesn't have a title.")
assert_is_not(my_boxplot.xaxis.get_label_text(), '', msg="Change the x-axis label to something more descriptive.")
assert_equal(my_boxplot.yaxis.get_label_text(), '', msg="The y-axis should not have a label.")
assert_equal(my_boxplot.get_data_ratio(),0.2681684097613303,
             msg="The function seems to be using the wrong ticker name")
```

# Problem 2:  Plot the distribution of days until next dividend

In this problem you will create a rugplot of days until next dividend for all the Dow Jones stocks in the dataset.

- Use the seaborn module to create the rugplot.


- Give the plot a descriptive title.


- Give the plot a descriptive x-axis label.


- Do *NOT* give the plot a y-axis label.

<!-- grade_id: p2-answer, points: 0, grade: false, solution: true -->
```python
def rug_plot(df):
    '''
    A function which creates a rugplot of the days until next dividend from the Dow Jones data.
        
    Parameters
    -------
    df: a pandas dataframe containing stock price data from the Dow Jones dataset    
    
    Returns
    --------
    
    ax: a matplotlib axes object containing a rugplot of the days to next dividend for all stocks in the 
        Dow Jones dataset    
    '''
        
    # YOUR CODE HERE
    
    
    return ax
```

```python
rp = rug_plot(df)
plt.show()
plt.close()
```

<!-- grade_id: p2-test, points: 1, grade: true, solution: false -->
```python
assert_is_instance(rp, mpl.axes.Axes, msg='Return a matplotlib axes object.')  
assert_is_not(len(rp.title.get_text()), 0, msg="Your plot doesn't have a title.")
assert_is_not(rp.xaxis.get_label_text(), '', msg="Change the x-axis label to something more descriptive.")
assert_equal(rp.yaxis.get_label_text(), '', msg="The y-axis should not have a label.")
assert_almost_equal(rp.get_data_ratio(),0.0027056, places=3,
            msg="The function is likely plotting the wrong data")
```

# Problem 3: Comparing Dow Jones index values


In this problem you will create a comparitive histogram to compare the distribution of the opening price of a stock and the distribution of the following weeks opening price of a stock.

- Create a 1-D comparative histogram (**side-by-side**) of the columns *open* and *next_weeks_open* from the Dow Jones dataset so that we can compare the current week's opening price to next week's opening price.

- Use default bin number(do not set bins).

- Create a label for the y-axis and the x-axis.

- Set a title for your plot.

- Create a legend so that we know which bars correspond to the current weeks open vs. next weeks open.




<!-- grade_id: p3-answer, points: 0, grade: false, solution: true -->
```python
def plot_compare_hist(df):
    '''
    A function which plots a comparative histogram to compare the current weeks opening price and 
    the next weeks opening price from the Dow Jones data.
        
    Parameters
    -------
    df: a pandas dataframe containing stock price data from the Dow Jones dataset
        
    Returns
    --------
    
    ax: a matplotlib axes object containing a comparative histogram of the open colum and next_weeks_open column from
        the Dow Jones dataset    
    '''
        
    # YOUR CODE HERE
    
    
    return ax
```

```python
ch = plot_compare_hist(df)
```

<!-- grade_id: p3-test, points: 1, grade: true, solution: false -->
```python
assert_is_instance(ch, mpl.axes.Axes, msg='Return an axes object.')  
assert_is_not(len(ch.title.get_text()), 0, msg="Your plot doesn't have a title.")
assert_is_not(ch.xaxis.get_label_text(), '', msg="Change the x-axis label to something more descriptive.")
assert_is_not(ch.yaxis.get_label_text(), '', msg="Change the y-axis label to something more descriptive.")
assert_is_instance(ch.legend_, mpl.legend.Legend, msg="Your plot doesn't have a Legend")
assert_equal(ch.patches[0].get_bbox().get_points()[0][0],12.1359,
             msg="The number of bins may not be correct, or the function is using the wrong ticker name")
assert_equal(ch.patches[0].get_bbox().get_points()[1][0],18.5995,
             msg="The number of bins may not be correct, or the function is using the wrong ticker name")
```

# Problem 4:  Plot horizontal barplot on average stock volume

- Use the seaborn barplot function to plot a **horizontal** bar plot on average stock volume. The y-axis should be stock symbols and the x-axis is average trading volume.

- Set figsize to (10, 8)

- Use the seaborn module to create the barplot

- Turn off confidence interval line

- Give the plot a descriptive title


<!-- grade_id: p4-answer, points: 0, grade: false, solution: true -->
```python
def plot_stock_volume(df):
    '''
    A  function to plot a horizontal bar plot on average stock volume. 
    
    Parameters
    -------
    df: a pandas dataframe containing stock price data from the Dow Jones dataset    
    
    Returns
    --------
    
    ax: a matplotlib axes object containing a barplot of the average stock volume    
    '''
    
    # YOUR CODE HERE
    
    
    return ax
```

```python
sv = plot_stock_volume(df)
plt.show()
plt.close()
```

<!-- grade_id: p4-test, points: 1, grade: true, solution: false -->
```python
assert_is_instance(sv, mpl.axes.Axes, msg='Return a matplotlib axes object.')  
assert_is_not(len(sv.title.get_text()), 0, msg="Your plot doesn't have a title.")
assert_almost_equal(sv.get_data_ratio(),3.9517929090328074e-08, places=3,
             msg="The function is plotting the wrong data or your plot is not a horizontal bar plot.")
```

# Problem 5: Comparing the current week's open to next week's open

Use matplotlib to plot scatter plot on open and next week's open for a particular stock.

- Create scatter plot on open and next week's open for a particular stock.

- The *open* variable should be on the x-axis.

- The *next week's open* variable should be on the y-axis.

- Create a label for the y-axis and the x-axis.

- Set a title for your plot


<!-- grade_id: p5-answer, points: 0, grade: false, solution: true -->
```python
def plot_scatter_plot(df, ticker_name):
    '''
    A function which plots scatter plot on open and next week's open for a particular stock.
        
    Parameters
    -------
    df: a pandas dataframe containing stock price data
    ticker_name: stock to plot
    
    Returns
    --------
    
    ax: a matplotlib axes object containing a scatter plot
    '''
        
    # YOUR CODE HERE
    
    
    return ax
```

```python
cb = plot_scatter_plot(df, 'IBM')
```

<!-- grade_id: p5-test, points: 1, grade: true, solution: false -->
```python
assert_is_instance(cb, mpl.axes.Axes, msg='Return a matplotlib axes object.')  
assert_is_not(len(cb.title.get_text()), 0, msg="Your plot doesn't have a title.")
assert_is_not(cb.xaxis.get_label_text(), '', msg="Change the x-axis label to something more descriptive.")
assert_is_not(cb.yaxis.get_label_text(), '', msg="Change the y-axis label to something more descriptive.")
assert_almost_equal(cb.get_data_ratio(),1, places=2,
             msg="The function is either plotting the wrong columns or wrong stock.")
```