# Module 2 Assignment


A few things you should keep in mind when working on assignments:

1. Run the first code cell to import modules needed by this assignment before proceeding to problems.
2. Make sure you fill in any place that says `# YOUR CODE HERE`. Do not write your answer anywhere else other than where it says `# YOUR CODE HERE`. Anything you write elsewhere will be removed or overwritten by the autograder.
3. Each problem has an autograder cell below the answer cell. Run the autograder cell to check your answer. If there's anything wrong in your answer, the autograder cell will display error messages.
4. Before you submit your assignment, make sure everything runs as expected. Go to the menubar, select Kernel, and Restart & Run all. If the notebook runs through the last code cell without error message, you've answered all problems correctly.
5. Make sure that you save your work (in the menubar, select File → Save and CheckPoint).

-----

# Run Me First!

```python
from nose.tools import assert_equal, assert_true
import warnings
warnings.filterwarnings("ignore", category=DeprecationWarning) 
```

# Problem 1: Create string variable

Create variable **s**, assign value "99.9" as a string to it.


<!-- grade_id: p1_ans, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
s = "99.9"
### END SOLUTION
```

<!-- grade_id: p1_test, points: 1, grade: true, solution: false -->
```python
assert_true(type(s)==str, msg="The type of s is not str")
assert_true(float(s)==99.9, msg="Value of s is not '99.9'")
```

# Problem 2: Calculate absolute value

- calcualte $97\times99-98^2$
- assign result to variable **result**
- assign absolute value of the result to variable **abs_result**


<!-- grade_id: p2_ans, points: 0, grade: false, solution: true -->
```python
result, abs_result = 0, 0
### BEGIN SOLUTION
result = 97*99-98**2
abs_result = abs(result)
### END SOLUTION
```

<!-- grade_id: p2_test, points: 1, grade: true, solution: false -->
```python
assert_equal(result, -1, msg="The result is wrong")
assert_equal(abs_result, 1, msg="The absolute result is wrong")
```

# Problem 3: Adding remainders

- Calculate the remainder of 7 divided by 6,    
   
   
- and calculate the remainder of 243 divided by 192. 


- Add both of these values together and save the result in a variable called "my_remainder".


<!-- grade_id: p3_ans, points: 0, grade: false, solution: true -->
```python
my_remainder = 0
### BEGIN SOLUTION
my_remainder = 7%6 + 243%192
### END SOLUTION
```

<!-- grade_id: p3_test, points: 1, grade: true, solution: false -->
```python
assert_equal(my_remainder,52,msg='my_remainder is not the correct value')

```

# Problem 4: Create a function to determine which value is greater  

- Write a function called find_greater that accepts 2 parameters called a and b.


- Inside of the function use an if statement to determine if a is greater than b.

    - If the value in a is greater than the value in b return variable a.
    
    - If the value in b is greater than the value in a return variable b.


<!-- grade_id: p4_ans, points: 0, grade: false, solution: true -->
```python
def find_greater(a, b):
    '''
    Returns the highest value between a and b.
    
    Parameters
    ----------
    a : an integer or floating point number
    b : an integer or floating point number
    
    Returns
    -------
    a floating point number that has the highest value
    '''
    ### BEGIN SOLUTION
    if a > b:
        return a
    else:
        return b
    ### END SOLUTION
```

<!-- grade_id: p4_test, points: 1, grade: true, solution: false -->
```python
assert_equal(find_greater(3, 4), 4, msg='The greater value between 3 and 4 is 4.')
assert_equal(find_greater(4, 3), 4, msg='The greater value between 4 and 3 is 4.')
assert_equal(find_greater(3.1459, 3.14), 3.1459, msg='The greater value between 3.1459 and 3.14 is 3.1459')

```

# Problem 5: Create a function to determine which value is the greatest

- Write a function called find_greatest that accepts 3 parameters called a, b, and c.


- Inside of the function use an if, else, and/or elif statements to determine which variable between a, b, and c has the highest value. Return this value.

**HINT: There's more than 1 way to solve this problem.**

<!-- grade_id: p5_ans, points: 0, grade: false, solution: true -->
```python
def find_greatest(a, b, c):
    '''
    Returns the highest value between a, b, and c.
    
    Parameters
    ----------
    a : an integer or floating point number
    b : an integer or floating point number
    c : an integer or floating point number

    Returns
    -------
    a floating point number that has the highest value
    '''
    ### BEGIN SOLUTION
    if a > b:
        if a > c:
            return a
    if b > a:
        if b > c:
            return b
    if c > a:
        if c > b:
            return c
    ### END SOLUTION
```

<!-- grade_id: p5_test, points: 1, grade: true, solution: false -->
```python
assert_equal(find_greatest(2, 3, 4), 4, msg='The greater value between 2, 3 and 4 is 4.')
assert_equal(find_greatest(3, 2, 4), 4, msg='The greater value between 3, 2 and 4 is 4.')
assert_equal(find_greatest(4, 3, 2), 4, msg='The greater value between 4, 3 and 2 is 4.')

```